# Lab 7 File & Function I

### <font color='red'>For debugging:

- You can try calling the function of interest with appropriate parameters.
- Understanding test cases are needed.
- Do not change the function names and the parameter names!!

### Exercise 1: Did you do it by yourself?
Write a Python function `check_exam(filenameA,filenameB)` that checks whether two student answer files are identical, indicating whether the student may have copied someone else's work. The function should compare the contents of two files, and return a boolean value:
 - **True**: if the contents of the two files are identical (suggesting potential copying).
 - **False**: if the contents are different.


**Example**
 - **stu1.txt**
```
A A B A D C C D A B
```
 - **stu5.txt**
```
A A B A D C C D A B
```

Calling `check_exam('stu1.txt','stu5.txt')` should return `True`.

**Example**
 - **stu2.txt**
```
C A B A D A C C A B
```
 - **stu5.txt**
```
A A B A D C C D A B
```

Calling `check_exam('stu2.txt','stu5.txt')` should return `False`.

<!-- grade_id: cell-0d462e84b395d749, points: 0, grade: false, solution: true -->
```python
def check_exam(filenameA, filenameB):
    # YOUR CODE HERE
    fileA = open(filenameA, 'r')
    fileB = open(filenameB, 'r')
    contentA = fileA.read()
    contentB = fileB.read()

    if contentA == contentB:
        return True
    else:
        return False
```

**Test `check_exam(filenameA,filenameB)`** <br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-0487c1cf3d359812, points: 1, grade: true, solution: false -->
```python
assert check_exam('stu1.txt','stu5.txt') == True
assert check_exam('stu1.txt','stu2.txt') == False
assert check_exam('stu2.txt','stu3.txt') == False
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-9295b820d98dbce5, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
```

<!-- grade_id: cell-6add70c9c5d3c76a, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
print('ok')
```

### Exercise 2: From Him to Her or Her to Him?
Write a Python function `check_letter(filename)` that receives the name of a **letterX.txt** file (e.g., letter1.txt) and returns who sent the letter to whom.

**Example**

**letter1.txt**

```
Dear Bella,
I miss you so much. The days seem longer, and the nights feel colder without you by my side.
I never meet whoever cute like you, with your beautiful smile and your warm heart.
You have this incredible ability to make everyone around you feel special, and I’m so lucky to have you in my life.

Sincerely, Sean
```

Calling `check_letter('letter1.txt')` should return a tuple of `('Bella','Sean')`.

<!-- grade_id: cell-61ca32f7f0b05d1f, points: 0, grade: false, solution: true -->
```python
def check_letter(filename) :
    # YOUR CODE HERE
    f = open(filename, "r" , encoding = "utf-8")
    sender = ""
    reciever = ""
    for x in f:
        if "Dear" in x:
            reciever = x
        if "Sincerely" in x:
            sender = x
    reciever = reciever.replace("Dear ", "")
    reciever = reciever.replace(",","")
    reciever = reciever.replace(" ","")
    reciever = reciever.replace("\n","")
    sender = sender.replace("Sincerely, ", "")
    sender = sender.replace(" ","")
    return sender,reciever
```

**Test `check_letter(filename)`** <br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-2fc0c9ee2915751d, points: 1, grade: true, solution: false -->
```python
assert check_letter('letter1.txt') == ('Sean','Bella')
assert check_letter('letter2.txt') == ('Siwakorn','Naruemon')
assert check_letter('letter3.txt') == ('CEO','Kafiling')
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-192223878ab324e2, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
print('ok')
```

Output:
```
ok
```

### Exercise 3: Mini Robot Movement Simulation
Write a Python function `move(filename,x,y)` that receives three parameters - a filename and an initial position of a robot as x and y coordinates (both integers). The function should read commands to move the robot based on the given file. The valid commands are one of the following per line:

`L`: Move left by one step \
`R`: Move right by one step \
`U`: Move up by one step \
`D`: Move down by one step

After processing all the commands, the function should return the final position of the robot (see the test cases).

**Example**

**move1.txt**
```
L
L
U
U
D
R
R
U
U
L
L
```

Calling `move('move1.txt',10,20)` should return a tuple of `(8,23)`.

**Example**

**move2.txt**
```
L
L
X
U
D
Z
R
```

Calling `move('move1.txt',10,20)` should return a tuple of `Invalid command`.

<!-- grade_id: cell-be8d7d9d314baa8c, points: 0, grade: false, solution: true -->
```python
def move(filename,x,y):
    # YOUR CODE HERE
    f = open(filename, 'r')
    for line in f:
        if "R" in line:
            x = x + 1
        elif "L" in line:
            x = x - 1
        elif "U" in line:
            y = y + 1
        elif "D" in line:
            y = y - 1
        else:
            return "Invalid command"
    return ("Robot stops at "+ str(x) + "," + str(y))
```

**Test `move(filename,x,y)`** <br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-dd3dea3eaeb0f42e, points: 1, grade: true, solution: false -->
```python
assert move('move1.txt',10,20) == 'Robot stops at 8,23' # ++
assert move('move1.txt',7,5) == 'Robot stops at 5,8' #++
assert move('move1.txt',100,-20) == 'Robot stops at 98,-17' #+-
assert move('move1.txt',-100,50) == 'Robot stops at -102,53' #-+
assert move('move1.txt',-100,-100) == 'Robot stops at -102,-97' #--
print("ok")
```

Output:
```
<ipython-input-8-f86b1cdcdecd> in <cell line: 1>()
----> 1 assert move('move1.txt',10,20) == 'Robot stops at 8,23' # ++
      2 assert move('move1.txt',7,5) == 'Robot stops at 5,8' #++
      3 assert move('move1.txt',100,-20) == 'Robot stops at 98,-17' #+-
      4 assert move('move1.txt',-100,50) == 'Robot stops at -102,53' #-+
      5 assert move('move1.txt',-100,-100) == 'Robot stops at -102,-97' #--

<ipython-input-7-e72428dc7587> in move(filename, x, y)
      1 def move(filename,x,y):
      2     # YOUR CODE HERE
----> 3     f = open(filename, 'r')
      4     for line in f:
      5         if "R" in line:

FileNotFoundError: [Errno 2] No such file or directory: 'move1.txt'
```

<!-- grade_id: cell-23108d5be0183034, points: 1, grade: true, solution: false -->
```python
assert move('move2.txt',15,20) == 'Invalid command' #
assert move('move2.txt',1,1) == 'Invalid command' #
print("ok")
```

<!-- grade_id: cell-4620d6aac4de8259, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
```

<!-- grade_id: cell-c060f2c780d90d0f, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
```

### Exercise 4: Print Lines   
Write two Python functions:
- `make_a_line(n)` that receives the length of a line (integer only) and returns a line consisting of that many asterisks (`*`).
- `make_3_lines(a,b,c)` that utilizes the function `make_a_line(n)` and return three lines of asterisks (`*`) according to the given parameters `a`,`b`,and `c`. See examples.

**Example**

Calling `make_3_lines(5,3,6)` should calls `make_a_line(5)`, `make_a_line(3)`, and `make_a_line(5)` and returns a string of

```
*****
***
******
```

**Example**

Calling `make_3_lines(10,10,10)` should calls `make_a_line(10)`, `make_a_line(10)`, and `make_a_line(10)` and returns a string of

```
**********
**********
**********
```

**Example**

Calling `make_3_lines(7,9,6)` should calls `make_a_line(7)`, `make_a_line(9)`, and `make_a_line(6)` and returns a string of

```
*******
*********
******
```

<!-- grade_id: cell-68b9df14aa3eb6b8, points: 0, grade: false, solution: true -->
```python
def make_a_line(n) :
    # YOUR CODE HERE
    return ("*"*n)
def make_3_lines(a,b,c):
    # YOUR CODE HERE
    make_a_line(a)
    make_a_line(b)
    make_a_line(c)
```

**Test `make_a_line(n)` and **`make_3_lines(a,b,c)`**<br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-5d81ac26485bb595, points: 1, grade: true, solution: false -->
```python
make_a_line(3) == "***"
make_a_line(5) == "*****"
make_a_line(2) == "**"
make_a_line(1) == "*"
make_a_line(0) == ""
print("ok")
```

<!-- grade_id: cell-5d81ac26485bb597, points: 1, grade: true, solution: false -->
```python
make_3_lines(3,5,2) == "***\n*****\n**"
make_3_lines(1,1,1) == "*\n*\n*"
make_3_lines(4,7,3) == "****\n*******\n***"
make_3_lines(0,0,0) == "\n\n"
make_3_lines(10,5,1) == "**********\n*****\n*"
print("ok")
```

<!-- grade_id: cell-e989a59583db6d2d, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
```

### Exercise 5: Summation
Write the following functions

1. `f(x)` This function should calculates and returns the value of $\frac{1} {x!}$ with eight decimal places.
2. `summation(a,b)` This function calculates and returns the summation of $f(x)$ when $x = a, a+1, ..., b$ so that is $\frac{1}{a!} + \frac{1}{(a+1)!} + ... + \frac{1}{(b-1)!} + \frac{1}{b!}$ with four decimal places.
The function should ensure that $a ≤ b$, otherwise return `0`.

**Tips**
- You can use function `factorial(n)` by import math module.
- You can use function `round(n)` to handle the number of decimal places.

<!-- grade_id: cell-42c0fa410330d920, points: 0, grade: false, solution: true -->
```python
import math

def f(x):
    # YOUR CODE HERE
    factorial = 1 / math.factorial(x)
    result = round(factorial, 8)
    return result

def summation(a, b):
    # YOUR CODE HERE
    if a > b:
        return 0
    total_sum = 0
    for i in range(a, b + 1):
        total_sum += f(i)
    return round(total_sum, 4)
```

**Test `f(x)` and `summation(a,b)`** <br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-cee1ece74b31fc86, points: 1, grade: true, solution: false -->
```python
assert f(0) == 1
assert f(1) == 1
assert f(2) == 0.5
assert f(7) == 0.00019841
print("ok")
```

<!-- grade_id: cell-98812d8b3fa5c74c, points: 1, grade: true, solution: false -->
```python
assert summation(1,3) == 1.6667
assert summation(3,6) == 0.2181
assert summation(5,6) == 0.0097
assert summation(3,4) == 0.2083
print("ok")
```

<!-- grade_id: cell-86ae967afc2f8808, points: 1, grade: true, solution: false -->
```python
'''Hidden Tests'''
```

### End of lab. Please remember to submit your work before the due date.
![Description of the image](https://s3.amazonaws.com/rails-camp-tutorials/blog/programming+memes/programming-or-googling.jpg)