# Lab 9 Collections 

### Exercise 1: Letter Checker
Write a fucntion `check_letter(phrase)` that receives a phrase or sentence and checks whether the phrase has all capital letters, all small letters, or a mix of both. 

<!-- grade_id: cell-5078ee1321ccc6da, points: 0, grade: false, solution: true -->
```python
def check_letter(phrase) :
    # YOUR CODE HERE
    if phrase.isupper():
        return "All capital letters."
    elif phrase.islower():
        return "All small letters."
    else:
        return "Mix letters."
```

**Test `check_letter`**   <br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-0e6439827a39f9c7, points: 1, grade: true, solution: false -->
```python
assert check_letter("EVERYTHING IS GOING TO BE ALRIGHT, JUST BELIEVE IN YOURSELF.") == "All capital letters."
assert check_letter("everything is going to be alright, just believe in yourself.") == "All small letters."
assert check_letter("Everything is going to be alright, just believe in yourself.") == "Mix letters."
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-192223878ab324e2, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 1'''
```

Output:
```
'Hidden Test 1'
```

<!-- grade_id: cell-643a1daa1a62c98b, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 2'''
```

Output:
```
'Hidden Test 2'
```

### Exercise 2: From One to Hundred 
Write a function `create_list(number)` for creating a list of all numbers from 1 to 100 that have the number of received parameter.

**Example** 
```
Enter a number between [0,9] : 0
[10, 20, 30, 40, 50, 60, 70, 80, 90, 100]
```

**Example** 
```
Enter a number between [0,9]: 3
[3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93]
```

**Example** 
```
Enter a number between [0,9]: 8
[8, 18, 28, 38, 48, 58, 68, 78, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 98]
```

<!-- grade_id: cell-f64231adf4d39779, points: 0, grade: false, solution: true -->
```python
def create_list(number): 
    # YOUR CODE HERE
    numbers = []
    for i in range(1, 101):
        if str(number) in str(i):
            numbers.append(i)
    return numbers
```

**Test `create_list`** <br>
Students, please run these cells below to check if your program works correctly.

<!-- grade_id: cell-7af2d59fd5e875e1, points: 1, grade: true, solution: false -->
```python
assert create_list(0) == [10, 20, 30, 40, 50, 60, 70, 80, 90, 100] 
assert create_list(3) == [3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93]
assert create_list(8) == [8, 18, 28, 38, 48, 58, 68, 78, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 98]
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-0487c1cf3d359812, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 1'''
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-6add70c9c5d3c76a, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 2'''
print('ok')
```

Output:
```
ok
```

### Exercise 3: Unique List 
Write a function called `unique_list` that receives a list and creates a new list with only the unique members from the original list. 

**Example**
```
Original List: 1,1,2,3,4,5,5,5,5,5,6,7,8,9
Unique List: 1,2,3,4,5,6,7,8,9
```

**Example**
```
Original List: a,g,b,a,v,c,c,c,b,g,y,k,y,k
Unique List: a,g,b,v,c,y,k
```

<!-- grade_id: cell-be8d7d9d314baa8c, points: 0, grade: false, solution: true -->
```python
def unique_list(original) :
    # YOUR CODE HERE
    list_1 = []
    unique_dict = {n:original.count(n) for n in original}
    for g in unique_dict.keys():
        list_1.append(g)
    
    return list_1
```

**Test `unique_list`** <br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-97585db5f1cdf60b, points: 1, grade: true, solution: false -->
```python
assert unique_list([1,1,2,3,4,5,5,5,5,5,6,7,8,9]) == [1,2,3,4,5,6,7,8,9]
assert unique_list([1.1,2.2,1.1,3.3,2.2,4.4,5.5]) == [1.1,2.2,3.3,4.4,5.5]
assert unique_list([10,20,20,10,30,40,10,20]) == [10,20,30,40]
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-7c5d23ec745720a5, points: 1, grade: true, solution: false -->
```python
assert unique_list(['a','a','a','a','a','a']) == ['a']
assert unique_list(['a','g','b','a','v','c','c','c','b','g','y','k','y','k']) == ['a','g','b','v','c','y','k']
assert unique_list(['apple','banana','apple','orange','banana','grape']) == ['apple','banana','orange','grape']
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-be895e6ef890eb46, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 1'''
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-c060f2c780d90d0f, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 2'''
print("ok")
```

Output:
```
ok
```

### Exercise 4: Charactor Pointing
Write a function `charactor_pointing(text,c)` that receives a string of text and an "interested" character as parameters. The function should return a new string where the **interested** character in the text is pointed to with a `^` symbol directly underneath it.

**Example**
```
Text: Welcome to Chulalongkorn University
Interested charactor: e
Welcome to Chulalongkorn University
 ^    ^                      ^     
```

<!-- grade_id: cell-68b9df14aa3eb6b8, points: 0, grade: false, solution: true -->
```python
def charactor_pointing(text,c) : 
    # YOUR CODE HERE
    c_location = []
    final_string = ""
    for letter in text:
        if c == letter:
            c_location.append("^")
        else:
            c_location.append(" ")
    for m in c_location:
        final_string = final_string + m
        
    return final_string
```

**Test `Charactor_pointing`**<br>
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-5d81ac26485bb595, points: 1, grade: true, solution: false -->
```python
assert charactor_pointing('Karfil Orange Cat','a').strip() == ' ^       ^     ^ '.strip()
assert charactor_pointing('Bachelor of Biotechnology','o').strip() == '      ^  ^    ^     ^ ^  '.strip()
assert charactor_pointing('Winnie-the-pooh','-').strip() == '      ^   ^    '.strip()
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-8f26126936a001be, points: 1, grade: true, solution: false -->
```python
assert charactor_pointing('MMMMMMMMMMMMMMMMMMMM','M').strip() == '^^^^^^^^^^^^^^^^^^^^'.strip()
assert charactor_pointing('MMMMMMMMMMMMMMMMMMMM','o').strip() == ''.strip()
assert charactor_pointing('MMMM','MM').strip() == ''.strip()
print("ok")
```

Output:
```
ok
```

<!-- grade_id: cell-198453a3256b2db5, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 1'''
### BEGIN HIDDEN TEST
assert charactor_pointing('Bell Warisara','a').strip() == '      ^   ^ ^'.strip()
assert charactor_pointing('Prof.Siwakorn','o').strip() == '  ^       ^  '.strip()
print('ok')
### END HIDDEN TEST
```

Output:
```
ok
```

<!-- grade_id: cell-5d81ac26485bb597, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 2'''
### BEGIN HIDDEN TEST
assert charactor_pointing('ppymph','p').strip() == '^^  ^ '.strip()
assert charactor_pointing('M I S S','I').strip() == '  ^    '.strip()
print('ok')
### END HIDDEN TEST
```

Output:
```
ok
```

### Exercise 5: Mixed String
Write three Python functions that each accept a string as a parameter. Each function should extract a specific type of character from the text and return it as a list:

 - `get_number(text)`: extracts all the numbers from the string.
 - `get_alphabet(text)`: extracts all the alphabetic characters from the string.
 - `get_symbol(text)`: extracts all the symbols (non-alphanumeric characters) from the string.

<!-- grade_id: cell-ebf76b7dba0dd0cb, points: 0, grade: false, solution: true -->
```python
def get_number(text):
    # YOUR CODE HERE
    results = []
    for m in text:
        if m.isdigit():
            results.append(int(m))
            
    return results
```

**Test `get_number(text)`**\
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-98812d8b3fa5c74c, points: 1, grade: true, solution: false -->
```python
assert get_number("M2er^r0y_Ch^24ristmas") == [2, 0, 2, 4]
assert get_number("abc123def456") == [1, 2, 3, 4, 5, 6]
assert get_number("no_numbers_here!") == []
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-746b7f42a9c49587, points: 1, grade: true, solution: false -->
```python
'''Hidden Test 1'''
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-21cc7663447ec940, points: 0, grade: false, solution: true -->
```python
def get_alphabet(text):
    # YOUR CODE HERE
    letters = []
    for char in text:
        if char.isalpha():
            letters.append(char)
    return letters
```

**Test `get_alphabet(text)`** \
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-2f48243cd74c0531, points: 1, grade: true, solution: false -->
```python
assert get_alphabet("M2er^r0y_Ch^24ristmas") == ['M', 'e', 'r', 'r', 'y', 'C', 'h', 'r', 'i', 's', 't', 'm', 'a', 's']
assert get_alphabet("123!@#$%^") == []
assert get_alphabet("abcXYZ") == ['a', 'b', 'c', 'X', 'Y', 'Z']
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-a243dbee84d08f53, points: 1, grade: true, solution: false -->
```python
'''Hidden Test'''
```

Output:
```
'Hidden Test'
```

<!-- grade_id: cell-aab9253c8a425d6b, points: 0, grade: false, solution: true -->
```python
def get_symbol(text):
    # YOUR CODE HERE
    symbols = []
    for char in text:
        if not char.isalnum():
            symbols.append(char)
    return symbols
```

**Test `get_symbol(text)`** \
Please run these cells below to check if your program works correctly.

<!-- grade_id: cell-506e0d48b3697c5c, points: 1, grade: true, solution: false -->
```python
assert get_symbol("M2er^r0y_Ch^24ristmas") == ['^', '_', '^']
assert get_symbol("123!@#$%^") == ['!', '@', '#', '$', '%', '^']
assert get_symbol("a!b@c#1*2^3") == ['!', '@', '#', '*', '^']
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-b151a0a0eb46ca10, points: 1, grade: true, solution: false -->
```python
'''Hidden Test'''
```

Output:
```
'Hidden Test'
```

### End of lab. Please remember to submit your work before due date.
![Description of the image](https://pleated-jeans.com/wp-content/uploads/2024/02/funny-programming-memes-february-11-2024-2.png)